## Задача 3 - Применение классификатора

Наконец, мы можем собрать воедино части решения, которые мы создали в проблемах 1 и 2. Здесь ваша цель - использовать
ваши новые функции и классифицировать данные по температуре в Фаренгейтах по четырем разным классам.

### Советы по выполнению упражнения

- Используйте **в точности** те же имена переменных, что и в инструкциях, потому что ваши ответы будут автоматически оценены, и тесты, которые оценивают ваши ответы, зависят от того же форматирования или именования переменных, что и в инструкциях.

- **Пожалуйста**,

    - **Не изменяйте имена файлов**. Сделайте все свое редактирование в файле `Exercise-4-зкщидуь-1.ipynb` (этом файле).
    - **Не копируйте/вставляйте ячейки в этом блокноте**. Мы используем автоматизированную систему оценки, которая не будет работать, если есть копии кодовых ячеек.
    - **Не изменяйте типы ячеек**. Вы можете добавить ячейки, но изменение типов ячеек для существующих ячеек (например, c кода на Markdown) также приведет к выходу из строя автоматизированной проверки.

### Условия успешного выполнения задания

- Создание нового **скрипта** с именем `temp_functions.py` с функциями, которые вы написали в задачах 1 и 2
- Наличие DocString в файле скрипта
- Преобразование и классификация списка температур в части 2
- Хранение значения температур в каждом классе в указанных переменных
- Ответ на вопросы в конце этой задачи
- Загрузка вашего блокнота в репозиторий

## Соглашение об использовании ИИ

**Введите свое имя в ячейку ниже** чтобы подтвердить, что вы следовали [принципам курса об использовании инструментов ИИ](https://geo-python-2024.readthedocs.io/ru/latest/course-info/ai-tools.html) и понимаете, что использование инструментов ИИ считается обманом.

<!-- grade_id: E4-P3-AI-LLM_statement, points: 0, grade: true, solution: true -->
Павел Федоров, это я

### Часть 1: Создание скрипта

Для выполнения этой части:

- Создайте новый **скрипт** c именем `temp_functions.py` в вашей директории, так же, как мы делали на уроке [Lesson 4 materials](https://geo-python-site.readthedocs.io/en/latest/notebooks/L4/script-files.html). Скопируйте и вставьте функции, которые вы написали в задачах 1 и 2 в этот файл сценария (`fahr_to_celsius` и `temp_classifier`)
- Добавьте DocString в файл скрипта (в дополнение к отдельным Docstrings в функциях)
- Сохраните файл скрипта и продолжайте работать в **этом** блокноте (`Exercise-4-problem-3.ipynb`)

### Часть 2: преобразование температуры из Фаренгейтов в Цельсии

Для этой части вы должны:

1. Запишите температуру Фаренгейта ниже в переменную `temp_data` (просто выполните ячейку ниже).

```python
# Список значений температуры (в градусах по Фаренгейту) в течение одной недели
temp_data =  [19, 21, 21, 21, 23, 23, 23, 21, 19, 21, 19, 21, 23, 27, 27, 28, 30, 30, 32, 32, 32, 32, 
              34, 34, 34, 36, 36, 36, 36, 36, 36, 34, 34, 34, 34, 34, 34, 32, 30, 30, 30, 28, 28, 27,
              27, 27, 23, 23, 21, 21, 21, 19, 19, 19, 18, 18, 21, 27, 28, 30, 32, 34, 36, 37, 37, 37, 
              39, 39, 39, 39, 39, 39, 41, 41, 41, 41, 41, 39, 39, 37, 37, 36, 36, 34, 34, 32, 30, 30,
              28, 27, 27, 25, 23, 23, 21, 21, 19, 19, 19, 18, 18, 18, 21, 25, 27, 28, 34, 34, 41, 37, 
              37, 39, 39, 39, 39, 41, 41, 39, 39, 39, 39, 39, 41, 39, 39, 39, 37, 36, 34, 32, 28, 28,
              27, 25, 25, 25, 23, 23, 23, 23, 21, 21, 21, 21, 19, 21, 19, 21, 21, 19, 21, 27, 28, 32,
              36, 36, 37, 39, 39, 39, 39, 39, 41, 41, 41, 41, 41, 41, 41, 41, 41, 39, 37, 36, 36, 34,
              32, 30, 28, 28, 27, 27, 25, 25, 23, 23, 23, 21, 21, 21, 19, 19, 19, 19, 19, 19, 21, 23,
              23, 23, 25, 27, 30, 36, 37, 37, 39, 39, 41, 41, 41, 39, 39, 41, 43, 43, 43, 43, 43, 43,
              43, 43, 43, 39, 37, 37, 37, 36, 36, 36, 36, 34, 32, 32, 32, 32, 30, 30, 28, 28, 28, 27,
              27, 27, 27, 25, 27, 27, 27, 28, 28, 28, 30, 32, 32, 32, 34, 34, 36, 36, 36, 37, 37, 37,
              37, 37, 37, 37, 37, 37, 36, 34, 30, 30, 27, 27, 25, 25, 23, 21, 21, 21, 21, 19, 19, 19,
              19, 19, 18, 18, 18, 18, 18, 19, 23, 27, 30, 32, 32, 32, 32, 32, 32, 34, 34, 34, 34, 34,
              36, 36, 36, 36, 36, 32, 32, 32, 32, 32, 32, 32, 32, 30, 30, 30, 30, 30, 30, 30, 30, 30,
              30, 30, 30, 30, 28, 28]
```

#### Часть 2 (продолжается)

2. Импортируйте функции `fahr_to_celsius` и` temp_classifier` из вашего скрипта `temp_functions.py`
3. Создайте пустой список под названием `temp_classes` (который будет заполнен номерами температурных классов позже)
4. **Преобразуйте температуру в Фаренгейтах** из списка `temp_data` в Цельсии и **клаццифицируйте температуры в Цельсиях** с помощью импортированных функций. Вывод должен храниться в списке `temp_classes`.
 
    **Вы можете сделать преобразование и классификацию температуры внутри одного цикла:**

    - Пройдите по списку температур в Фаренгейтах, то есть, создайте цикл `for` и внутри цикла:

        1. Создайте новую переменную под названием `temp_celsius`, в которой вы будете хранить температуру в Цельсиях, используя функцию `fahr_to_celsius` для преобразования температуры
        2. Создайте новую переменную с именем `temp_class`, в которой вы назначаете номер класса температуры (0, 1, 2 или 3) с помощью функции ` temp_classifier`
        3. Добавьте значение `temp_class` в список `temp_classes`

<!-- grade_id: problem_3_ag_p2_functions_usage, points: 0, grade: false, solution: true -->
```python
from temp_functions import fahr_to_celsius, temp_classifier
temp_classes = []

for i in temp_data:
    temp_celsius = fahr_to_celsius(i) 
    temp_class = temp_classifier(temp_celsius) 
    temp_classes.append(temp_class)
```

Output:
```
ModuleNotFoundError                       Traceback (most recent call last)
Cell In[30], line 1
----> 1 from temp_functions import fahr_to_celsius, temp_classifier
      2 temp_classes = []
      4 for i in temp_data:

ModuleNotFoundError: No module named 'temp_functions'
```

#### Часть 2

6. Рассчитайте, сколько температур в каждом температурном классе:

    1. Создайте четыре переменных с именами `zeros`, `ones`, `twos`, и `threes` 
    2. Посчитайте сколько раз значения 0, 1, 2 и 3 присутствуют в списке `temp_class` и распечатайте результаты ниже. 
  
**Подсказка**: Вы можете использовать [функцию **count()**](https://www.tutorialspoint.com/python3/list_count.htm) или цикл `for`.

```python
# 1. Сколько значений 0 существует в списке temp_class?
zeros=temp_classes.count(0)
print(zeros)
# 2. Сколько значений 1 существует в списке temp_class?
ones=temp_classes.count(1)
print(ones)
# 3. Сколько значений 2 существует в списке temp_class?
twos=temp_classes.count(2)
print(twos)
# 4. Сколько значений 3 существует в списке temp_class?
threes=temp_classes.count(3)
print(threes)
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[20], line 2
      1 # 1. Сколько значений 0 существует в списке temp_class?
----> 2 zeros=temp_classes.count(0)
      3 # 2. Сколько значений 1 существует в списке temp_class?
      4 ones=temp_classes.count(1)

NameError: name 'temp_classes' is not defined
```

### Задача 3 - тесты

Запустите ячейки ниже, чтобы проверить свой ответ. Обратите внимание, что вы не можете изменять тестовые ячейки. Ваше решение верно, если вы не видите никаких ошибок при запуске этих ячеек.

<!-- grade_id: problem_3_ag_p2_functions_usage_test1, points: 0.5, grade: true, solution: false -->
```python
import inspect

# Убедитесь, что функции находятся в пространстве имен
assert inspect.isfunction(fahr_to_celsius)
assert inspect.isfunction(temp_classifier)
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[21], line 4
      1 import inspect
      3 # Убедитесь, что функции находятся в пространстве имен
----> 4 assert inspect.isfunction(fahr_to_celsius)
      5 assert inspect.isfunction(temp_classifier)

NameError: name 'fahr_to_celsius' is not defined
```

<!-- grade_id: problem_3_ag_p2_functions_usage_test2, points: 0.5, grade: true, solution: false -->
```python
# Убедитесь, что переменная была создана
assert 'temp_celsius' in locals()
assert 'temp_class' in locals()
assert 'temp_classes' in locals()

# Убедитесь, что temp_classes - это список
assert type(temp_classes) == list
```

Output:
```
AssertionError                            Traceback (most recent call last)
Cell In[22], line 2
      1 # Убедитесь, что переменная была создана
----> 2 assert 'temp_celsius' in locals()
      3 assert 'temp_class' in locals()
      4 assert 'temp_classes' in locals()

AssertionError:
```

<!-- grade_id: problem_3_ag_p2_functions_usage_test3, points: 0.5, grade: true, solution: false -->
```python
# Убедитесь, что функции имеют один параметр
t_params = list(inspect.signature(temp_classifier).parameters.keys())
f_params = list(inspect.signature(fahr_to_celsius).parameters.keys())
assert len(t_params) == 1
assert len(f_params) == 1
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[23], line 2
      1 # Убедитесь, что функции имеют один параметр
----> 2 t_params = list(inspect.signature(temp_classifier).parameters.keys())
      3 f_params = list(inspect.signature(fahr_to_celsius).parameters.keys())
      4 assert len(t_params) == 1

NameError: name 'temp_classifier' is not defined
```

Проверьте, что требуемые переменные существуют, и распечатайте их значение (проверьте вручную, что ответы имеют смысл!):

<!-- grade_id: problem_3_ag_p2_functions_usage_test4, points: 0.5, grade: true, solution: false -->
```python
# Проверьте переменную "zeros"
assert 'zeros' in locals()
print(zeros)
```

Output:
```
AssertionError                            Traceback (most recent call last)
Cell In[24], line 2
      1 # Проверьте переменную "zeros"
----> 2 assert 'zeros' in locals()
      3 print(zeros)

AssertionError:
```

<!-- grade_id: problem_3_ag_p2_functions_usage_test5, points: 0.5, grade: true, solution: false -->
```python
# Проверьте переменную "ones" 
assert 'ones' in locals()
print(ones)
```

Output:
```
AssertionError                            Traceback (most recent call last)
Cell In[25], line 2
      1 # Проверьте переменную "ones" 
----> 2 assert 'ones' in locals()
      3 print(ones)

AssertionError:
```

<!-- grade_id: problem_3_ag_p2_functions_usage_test6, points: 0.5, grade: true, solution: false -->
```python
# Проверьте переменную "twos" 
assert 'twos' in locals()
print(twos)
```

Output:
```
AssertionError                            Traceback (most recent call last)
Cell In[26], line 2
      1 # Проверьте переменную "twos" 
----> 2 assert 'twos' in locals()
      3 print(twos)

AssertionError:
```

<!-- grade_id: problem_3_ag_p2_functions_usage_test7, points: 0.5, grade: true, solution: false -->
```python
# Проверьте переменную "threes" 
assert 'threes' in locals()
print(threes)
```

Output:
```
AssertionError                            Traceback (most recent call last)
Cell In[27], line 2
      1 # Проверьте переменную "threes" 
----> 2 assert 'threes' in locals()
      3 print(threes)

AssertionError:
```

### Часть 3: Проверьте свой код
 
- Не забудьте проверить, что вы включили комментарии в свой код и использовали требуемые имена переменных
- Не забудьте сделать **commit и push** свой скрипта `temp_functions.py` в свой собственный репозиторий для упражнения 4!
- Если у вас есть какие-либо мысли и проблемы на этом этапе, вы можете написать их ниже.

<!-- grade_id: problem_3_mg_p3_check_your_code, points: 0.5, grade: true, solution: true -->
Вообще не видит скрипт, который сохраняю, постоянно выдает ошибку и пишет, что нет такого скрипта(я в юпитерлайт всё делал, может в нем проблема). Сделал в replit без скрипта, а просто с функциями, там всё получилось, поэтому чтобы хоть что-то прислать, просто вначале добавил, что якобы сделал импорт из скрипта

#### Все сделано!

Вот и все!Теперь вы закончили все проблемы для упражнения 4! Не забудьте запушить все свои изменения в ваш репозиторий (**включая скрипт `temp_functions.py`**).